# E1 : Data Wrangling

In this week's workbook, you'll work through an example in this to gain practice on how to:
* read a DataFrame into Python
* wrangle the data
* clean the data

# Part I: Setup

Data wrangling often requires additional functionality outside what's included in Python by default. For this, we'll import other functions from helpful packages.

**Import the following packages using their common shortened name found in parentheses:**

* `numpy` (`np`)
* `pandas` (`pd`)

<!-- grade_id: cell-b689d21f629d9dda, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE



```

```python
# This is a test cell!  If you run this and there are no errors you probably did it right!
assert pd
assert np
```

**Run the following cell code to make things throughout the rest of this workbook a little prettier.** (Note: You don't have to edit code here, but are free to and see what changes to be sure you understand each line.)

```python
# Configure libraries
# there are so many settings... you should probably explore the documentation and see what they are!

# method 1 - directly setting values
# Don't display too many rows/cols of DataFrames
pd.options.display.max_rows = 7
pd.options.display.max_columns = 8

# method 2 - using set_option
# Round decimals when displaying DataFrames
pd.set_option('display.precision', 2)
```

**Read the CSV file at the URL https://raw.githubusercontent.com/fivethirtyeight/data/master/steak-survey/steak-risk-survey.csv into Python in and assign it to the variable `survey`**.

<!-- grade_id: cell-65e31c6985ba9dd2, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE





```

<!-- grade_id: cell-0fdae0bebda7c436, points: 0, grade: true, solution: false -->
```python
# This is a test cell!  If you run this and there are no errors you probably did it right!
assert survey.shape == (551, 15)
```

These data contain survey responses from Americans who responded to a SurveyMonkey Audience poll. These data were used in the [FiveThirtyEight](https://fivethirtyeight.com) article: *[How Americans Like Their Steak](https://fivethirtyeight.com/features/how-americans-like-their-steak/)*


Now might be a good time to explore the data and see what it looks like... 

# Part II: Wrangling

**Write a line of code to look at the first few rows of the DataFrame**

<!-- grade_id: cell-c7cd7d4cfbf98114, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE





```

What do you notice about the first row of the dataframe? Notice that it's not actually an observation from a respondent? **Remove this row from the dataset. Assign this back to the variable `survey`. Print the first few rows again to make sure you've accomplished this.**

<!-- grade_id: cell-0624532e0d0d4ce1, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE






```

<!-- grade_id: cell-3969e31aae66b836, points: 0, grade: true, solution: false -->
```python
# This is a test cell!  If you run this and there are no errors you probably did it right!
assert survey.shape == (550, 15)
```

Notice that there are a lot of different questions that were asked of respondents (columns) and 550 people who responded to the survey (rows). 

**Print a list of all the column names in this DataFrame.**

We'll only end up working with a subset of these.

<!-- grade_id: cell-c91f7a1c6729ad43, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE






```

Now we have a sense of what information is included in the dataset. In the coming weeks, we'll answer the following questions:
1. Who cheats more on their significant other - males or females?
2. Are cigarette smokers less likely to skydive?
3. Do people in New England gamble more than other parts of the country?

To answer these we'll only need data from *some* of the columns in the dataset.

Let's drop the columns we don't need. **Drop the first two columns from the dataset. This should still be assigned to the variable `survey`.**

<!-- grade_id: cell-8a1106312081bab5, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE






```

<!-- grade_id: cell-89822f596640cd4e, points: 0, grade: true, solution: false -->
```python
# This is a test cell!  If you run this and there are no errors you probably did it right!
assert survey.shape == (550,13)
```

Now that we've got the columns we want, let's clean up those column names. **Rename the columns in `survey` so the appropriate columns have the following names:**

* smoking
* alcohol
* gambling
* skydiving
* speeding
* cheated
* steak
* steak_preference
* gender
* age
* income 
* education
* region

Note that a list has been provided here, with these names, so you don't have to type them all:

```python
['smoking','alcohol','gambling',
 'skydiving','speeding', 'cheated',
 'steak', 'steak_preference','gender',
 'age', 'income', 'education', 'region']
```

<!-- grade_id: cell-ff9ae6330821a062, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE






```

<!-- grade_id: cell-ac6bcfa7d1ac939d, points: 0, grade: true, solution: false -->
```python
# This is a test cell!  If you run this and there are no errors you probably did it right!
assert list(survey) == ['smoking',
                        'alcohol',
                        'gambling',
                        'skydiving',
                        'speeding',
                        'cheated',
                        'steak',
                        'steak_preference',
                        'gender',
                        'age',
                        'income',
                        'education',
                        'region']
```

We're in pretty good shape now. **Print the first few rows of the `survey` DataFrame to see what the data look like at this point.**

```python
### YOUR CODE HERE






```

# Part III: Cleaning

**Now that we've go the data we need, let's get a sense for how much missing data there is in this dataset by determining how many null-containing rows there are in `survey`. Assign this value to the variable `null_rows`.**

<!-- grade_id: cell-5bc6a10438afe331, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE






```

<!-- grade_id: cell-e43974ae43ee28d8, points: 0, grade: true, solution: false -->
```python
# This is a test cell!  If you run this and there are no errors you probably did it right!
assert null_rows == 217
```

Good to know that lots of people didn't answer every question. We'll keep that in mind as we work with this dataset.

Simply dropping missing observations is typically not good practice; however, in this case we'll drop observations that have missing data across the entire row, as these are individuals who didn't participate in the survey at all. **Remove rows where ALL the columns have missing data for that participant.**

```python
# here's a hint
survey.dropna?
```

<!-- grade_id: cell-b8417c0e50579e72, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE






```

<!-- grade_id: cell-27d421b2b5aaac0f, points: 0, grade: true, solution: false -->
```python
# This is a test cell!  If you run this and there are no errors you probably did it right!
assert survey.shape == (541, 13)
```

**Print the first few rows to remind yourself what the data look like at this point.**

<!-- grade_id: cell-e1148cbe6b8da30b, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE






```

Note that the first row no longer has all missing data here. We've got a dataset we can work with now!

**Great work on this workbook! We'll continue to work with this dataset in section to answer our questions of interest. You can work on other parts of the class or help a classmate work through this - we always understand things best once we've had to explain them to someone else.**